# a4 - Python

This assignment will cover topics of text mining and clustering

Make sure that you keep this notebook named as "a4.ipynb" 

Any other packages or tools, outside those listed in the assignments or Canvas, should be cleared
by Dr. Brown before use in your submission.

# Q0 - Setup

The following code looks to see whether your notebook is run on Gradescope (GS), Colab (COLAB), or the linux Python environment you were asked to setup.

```python
import re 
import os
import platform 
import sys 

# flag if notebook is running on Gradescope 
if re.search(r'am', platform.uname().release): 
    GS = True
else: 
    GS = False

# flag if notebook is running on Colaboratory 
try:
  import google.colab
  COLAB = True
except:
  COLAB = False

# flag if running on Linux lab machines. 
cname = platform.uname().node
if re.search(r'(guardian|colossus|c28|coc-15954-m)', cname):
    LLM = True 
else: 
    LLM = False

print("System: GS - %s, COLAB - %s, LLM - %s" % (GS, COLAB, LLM))
```

Output:
```
System: GS - False, COLAB - False, LLM - True
```

## Notebook Setup

It is good practice to list all imports needed at the top of the notebook. You can import modules in later cells as needed, but listing them at the top clearly shows all which are needed to be available / installed.

If you are doing development on Colab, the otter-grader package is not available, so you will need to install it with pip (uncomment the cell directly below).

```python
# Only uncomment if you developing on Colab 
# if COLAB == True: 
#     print("Installing otter:")
#     !pip install otter-grader==4.2.0 
```

```python
# Import standard DS packages 
import pandas as pd
import numpy as np
import matplotlib as mpl
import matplotlib.pyplot as plt
import seaborn as sns
import math
import scipy
import statistics
import textwrap
%matplotlib inline


from sklearn.model_selection import train_test_split, StratifiedKFold 
from sklearn.preprocessing import StandardScaler, MinMaxScaler
from sklearn.pipeline import Pipeline, make_pipeline
from sklearn.model_selection import GridSearchCV
from sklearn import tree        # decision tree classifier
from sklearn import neighbors   # knn classifier
from sklearn import naive_bayes # naive bayes classifier 
from sklearn import svm         # svm classifier
from sklearn import ensemble    # ensemble classifiers
from sklearn import metrics     # performance evaluation metrics
from sklearn import model_selection
from sklearn import preprocessing 
from sklearn.decomposition import PCA
from sklearn.datasets import load_files
from sklearn.feature_extraction.text import CountVectorizer
from sklearn.feature_extraction.text import TfidfVectorizer, TfidfTransformer

from sklearn import preprocessing
from sklearn.cluster import AgglomerativeClustering
from scipy.spatial.distance import pdist
from scipy.spatial.distance import squareform
from sklearn import cluster
from scipy.cluster.hierarchy import dendrogram
from scipy.cluster import hierarchy

# Package for Autograder 
import otter 
grader = otter.Notebook()
```

```python
grader.check("q0")
```

Output:
```
q0 results: All test cases passed!
```

# Q1 - Text Classification 


You will look to predict whether scenes in Shakespeare's plays come from the comedies or histories.  Shakespeare's comedies include plays such as: The Taming of the Shrew, The Merchant of Venice, Much Ado About Nothing, and more.  The histories include: Richard II, Richard III, Henry IV part 1, Henry IV part 2, Henry V, Henry VI (part 1-3). 

The plays were downloaded from the [Shakespeare Corpus](http://hdl.handle.net/11040/24448).  Note, the original plays were downloaded from [Project Gutenberg](https://www.gutenberg.org/). 

Note, the plays have already had significant preprocessing.  The plays have been scrubbed by: removing digits, making the file all lowercase, and removing punctuation, excluding hyphens and word-internal apostrophes. Also, the character names and stage directions have been removed manually. An example of the text would be like this:

*Before scrubbing:*

    ADAM. Yonder comes my master, your brother.
    ORLANDO. Go apart, Adam, and thou shalt hear how he will shake me
    up. [ADAM retires]
    OLIVER. Now, sir! what make you here?

*After scrubbing:*

    yonder comes my master your brother
    go apart adam and thou shalt hear how he will shake me
    up
    now sir what make you here

The text files are split into negative - comedies and positive - histories.  


## Q1(a) - Load the Data 

Load the plays into a list `textdata` and a np.ndarray `yvalues`.  I highly suggest using `scikit-learn`'s `load_files` function, with the `random_state` set to 42.  

 

```python
# Load the plays data     
plays = load_files("data/shakespeare", random_state=42)

# Extract text data
textdata = dataset.data

# Extract target values
yvalues = dataset.target

print("Samples per class: {}".format(np.bincount(yvalues)))

plays.filenames[0:10]
```

Output:
```
File ~/.conda/envs/cs4821/lib/python3.10/pathlib.py:1126, in Path.read_bytes(self)
   1122 def read_bytes(self):
   1123     """
   1124     Open the file in bytes mode, read it, and close the file.
   1125     """
-> 1126     with self.open(mode='rb') as f:
   1127         return f.read()

File ~/.conda/envs/cs4821/lib/python3.10/pathlib.py:1119, in Path.open(self, mode, buffering, encoding, errors, newline)
   1117 if "b" not in mode:
   1118     encoding = io.text_encoding(encoding)
-> 1119 return self._accessor.open(self, mode, buffering, encoding, errors,
   1120                            newline)

IsADirectoryError: [Errno 21] Is a directory: 'data/shakespeare/histories/.ipynb_checkpoints'
```

```python
grader.check("q1a")
```

## Q1(b) - Prepare the Data 

Split the data into `text_trainval`, `text_test` and `y_trainval`, `y_test` variables.  Use 20% of the data in the test set with a `random_state` of 42 and make sure to stratify the split (the data is imbalanced). 

```python
# Split the data 



```

```python
grader.check("q1b")
```

## Q1(c)  - Explore the Data

Create a document-term count matrix for the "trainval" data using the default tokenizer, removing the standard English stopwords and store this in `dtm_trainval`.

Store the names of the terms in the dtm matrix in the variable `vocab`.

```python
# Create document-term count matrix for the "trainval" text data 



```

```python
grader.check("q1c")
```

<!-- BEGIN QUESTION -->

## Q1(d) - Explore the Data 

Create a plot showing the top 15 most frequently used words in the trainval text data. 

```python
# Create a plot of the top 15 most frequently used words 



```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

## Q1(e) - Explore the Data 

For the trainval text data, plot the top 15 most frequently used words in the histories and the comedies.  Put these two bar plots side-by-side to compare the results. 

```python
# Create a plot of the top 15 most frequently used words in the 
#  Comedies and Histories. 



```

<!-- END QUESTION -->

## Q1(f) - Bernoulli Naive Bayes 

Let's know explore using Bernoulli Naive Bayes as a classifier, `bern_nb`, to predict the type of play. 

We will use the split of the data into trainval / test found above to train the model and then evaluate it's performance. 

Create the training data, `X_trainval` to be binary with features using the default tokenizer, stop words removed, appear in at least 5 documents and is limited to the top 5000 features.  

Calculate the training accuracy `train_acc_bern` and testing accuracy `test_acc_bern` for the model.



```python
# Run Bernoulli Naive Bayes model



```

```python
grader.check("q1f")
```

## Q1(g) - Multinomial Naive Bayes 

Let's know explore using multinomial Naive Bayes as a classifier, `mult_nb`, to predict the type of play. 

We will use the split of the data into trainval / test found above to train the model and then evaluate it's performance. 

Create the training data, `X_trainval` with features using the default tokenizer, stop words removed, appear in at least 5 documents and is limited to the top 5000 features.  

Calculate the training accuracy `train_acc_mult` and testing accuracy `test_acc_mult` for the model.


```python
# Run Multinomial Naive Bayes model



```

```python
grader.check("q1g")
```

<!-- BEGIN QUESTION -->

## Q1(h) - Naive Bayes Models 

Looking at the results of the two models above. Answer the following questions.  


Which of the two models is preferred?  Why?   (10 words or less)


What is a problem for both models?  How might you solve it? (12 words or less)

<!-- END QUESTION -->

## Q1(i) - Other Models 

Let's now look to explore using other models. 

You will set up a pipeline, `pipe`, that will use a Random Forest model with 100 trees and a `random_state` = 42.  

In the pipeline (`param_grid`), you will consider using both a document term count matrix as well as a TF-IDF matrix. 
In either case, limit the matrix to words that appear in at least 5 documents and remove English stop words.  Consider features of unigrams, unigrams + bigrams, and bigrams.  Examine a maximum feature limit of either 2500 or 5000.  

Optimize your choice of hyperparameters using GridSearchCV, `grid`, with stratified 5-fold cross-validation (random_state = 42), select the parameters using AUC. (See how to set up the scorer below)

Note, do not run the jobs in parallel, you may exceed the memory resources of the autograder on Gradescope. 

```python
# Run Pipeline to find best model

pipe = ...

param_grid = ...

cvStrat = ...

score_fn = metrics.make_scorer(metrics.roc_auc_score, needs_threshold = False)
grid = ...


print("Best cross-validation score: {:.2f}".format(grid.best_score_))
print("Best parameters:\n{}".format(grid.best_params_))
```

```python
grader.check("q1i")
```

## Q1(j) - Explore the Results 

Calculate the AUC on the test text, `auc_test`.   

Gather the importances of the features in the best model in `importance`. [Feature Importance Example](https://scikit-learn.org/stable/auto_examples/ensemble/plot_forest_importances.html)

Create a bar plot with the top 10 features sorted by importance.  


```python
# Calculate performance on test data `auc_test` 

# Create plot of top 10 features sorted by importance. 



```

```python
grader.check("q1j")
```

# Q2

Consider methods to cluster NBA players based on their statistics. 




## Q2(a) 

Load in data for NBA players from the 2018-2019 season. 

Filter the players to only consider those who have played in more than 20 games.  

Ignore the first 7 columns as well as ignore columns of statistics with percentages (FG%, 3P%, 2P%, eFG%, FT%). 

```python
# Load in data and filter out requested rows and columns. 

nba = ...



nba.head()
```

```python
grader.check("q2a")
```

## Q2(b)

The features have different ranges, therefore we should scale the data before considering the clustering analysis. Scale the data using min-max normalization with range of [0, 1].

```python
# Scale the data 

scaler = ...
nbaScaled = ...

```

```python
grader.check("q2b")
```

## Q2(c)

Run Kmeans clustering on the data with k=2, . . . , 12.  Set the `random_state` in the Kmeans method to 42 and `n_init` to 10. 
For each value of k, keep track of the within-cluster variation (this quantity is referred to as different terms such as “inertia” and total “within-cluster sum-of-squares”), the Calinski-Harabasz score, and the Davies-Bouldin index on the resulting clusters. 

```python
# Run Kmeans 

sse = []
dbscore = []
chscore = []



```

```python
grader.check("q2c")
```

## Q2(d)

Assuming the best number of clusters is 4 (depending on which measure we use different number of clusters is preferred with this data). 

Run Kmeans again with this value for $k$ (use `n_init` = 10 and `random_state` = 10).

Create a DataFrame `clusterStats` with the mean statistics (centers) of each group. 

The DataFrame should have rows for each cluster group 0, 1, 2, 3 and columns for the mean statistics.  

Add a column `Num` reporting the number of samples in each group. 


```python
# Create a Data Frame for the mean statistics of each group 


clusterStats[['Num', 'MP', 'FG', '3P', 'FT']]
```

```python
grader.check("q2d")
```

## Q2(e)

Report the same statistics as in (e), but using the original data scaling (reverse the scaling back to the original data range). 

Store results in `clusterStatsOrig`; this DataFrame should not have the "Num" column.



```python
# Create a Data Frame for the mean statistics of each group (using the 
#   original data scaling)


clusterStatsOrig = ...

clusterStatsOrig[['MP', 'FG', '3P', 'FT']]
```

```python
grader.check("q2e")
```

<!-- BEGIN QUESTION -->

## Q2(f) 

Apply PCA to the basketball data.  Plot the first two principal components, colored by the best group labels found above.  

```python
# Run PCA on the nba data and plot the first two principal components
#  colored by the group labels. 



```

<!-- END QUESTION -->

# Q3 : Clustering - Spotify Music

For this problem you will look at popular streaming music.  Specifically, Spotify's top 100 streaming songs.  For each song information about the song is described with different properties: `duration`, `energy`, `key`, etc. 


## Q3(a) - Load and Prepare the Data 

Load in the `music.csv` data.  

The clustering algorithms will only consider variables of `duration` to the end of the DataFrame. 

Standardize the variables to be used in clustering.  

```python
# Load in music data 

music = pd.read_csv(...)

music.head()
```

```python
grader.check("q3a")
```

## Q3(b) - Hierarchical Clustering 

Perform Hierarchical clustering with **single** linkage on just the top 30 songs. 

Report results in a dendrogram, `dg_single` and label the samples by the Artist.  



```python
# Perform Hierarchical clustering with single linkage on top 30 songs 
# Report results in a dendrogram, dg_single

```

```python
grader.check("q3b")
```

## Q3(c) - Hierarchical Clustering, part 2 

Perform Hierarchical clustering with **complete** linkage on just the top 30 songs.

Report results in a dendrogram, `dn_complete` and label the samples by the Artist.

```python
# Perform Hierarchical clustering with complete linkage on top 30 songs 
# Report results in a dendrogram, dg_complete

```

```python
grader.check("q3c")
```

## Q3(d) - Hierarchical Clustering, part 3

Perform Hierarchical clustering with **aveage** linkage on just the top 30 songs.

Report results in a dendrogram, `dn_average` and label the samples by the Artist.

```python
# Perform Hierarchical clustering with average linkage on top 30 songs 
# Report results in a dendrogram, dg_average

```

```python
grader.check("q3d")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

**NOTE** the submission must be run on the campus linux machines.  See the instruction in the Canvas assignment.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q0 (points: 2)
```
>>> COLAB == False
True
```
```
>>> (LLM == True) | (GS == True)
True
```


Tests for q1a (points: [2, 2, 2, 2, 2, 2])
```
>>> len(textdata) == 327
True
```
```
>>> all(np.bincount(yvalues) == [119, 208])
True
```
```
>>> textdata[0].decode()[0:40] == '\r\ni muse my lord of gloucester is not co'
True
```
```
>>> textdata[10].decode()[0:40] == '\r\nkill the poys and the luggage tis expr'
True
```
```
>>> textdata[23].decode()[200:250] == ' bishop here at whose hands\r\nhe hath good usage an'
True
```
```
>>> textdata[13].decode()[50:100] == 'ir give me your hand sir an early stirrer by\r\nthe '
True
```


Tests for q1b (points: [2, 2, 2, 2, 1, 1])
```
>>> len(text_trainval) == 261
True
```
```
>>> len(text_test) == 66
True
```
```
>>> all(y_trainval[0:10] == [1, 0, 1, 0, 0, 1, 1, 0, 0, 1])
True
```
```
>>> all(y_test[0:10] == [0, 0, 1, 0, 1, 1, 1, 1, 1, 1])
True
```
```
>>> text_trainval[0].decode()[500:530] == 'irds\r\nconceive when after many'
True
```
```
>>> text_test[0].decode()[1000:1040] == 'er than those that are so washed how muc'
True
```


Tests for q1c (points: 6)
```
>>> dtm_trainval.shape == (261, 13953)
True
```
```
>>> len(vocab) == 13953
True
```
```
>>> type(dtm_trainval) == scipy.sparse._csr.csr_matrix
True
```
```
>>> all(vocab[450:460] == ['answer', 'answerable', 'answered', 'answers', 'anthony', 'antic', 'anticly', 'antics', 'antiopa', 'antipholus'])
True
```


Tests for q1f (points: 8)
```
>>> np.isclose(train_acc_bern, 0.95019157)
True
```
```
>>> all(X_trainval[0, 1:10].toarray()[0] == [0, 0, 0, 0, 0, 0, 0, 0, 0])
True
```
```
>>> all(X_trainval[5, 1280:1290].toarray()[0] == [0, 1, 0, 0, 0, 0, 0, 0, 0, 1])
True
```


Tests for q1g (points: 8)
```
>>> train_acc_mult == 1.0
True
```
```
>>> all(X_trainval[0, 1:10].toarray()[0] == [0, 0, 0, 0, 0, 0, 0, 0, 0])
True
```
```
>>> all(X_trainval[5, 1280:1290].toarray()[0] == [0, 2, 0, 0, 0, 0, 0, 0, 0, 1])
True
```


Tests for q1i (points: [3])
```
>>> type(grid.best_params_['vec']) == TfidfVectorizer
True
```


Tests for q1j (points: 8)
```
>>> all(np.isclose(importance[0:10], [0.0, 0.0, 0.00054952, 0.00035462, 0.00024037, 0.0, 0.00077599, 0.00040087, 0.0, 0.0]))
True
```


Tests for q2a (points: [2.5, 0.5, 0.5, 0.5, 0.5, 0.5, 2, 2])
```
>>> nba.shape == (420, 25)
True
```
```
>>> not 'FG%' in nba.columns
True
```
```
>>> not '3P%' in nba.columns
True
```
```
>>> not '2P%' in nba.columns
True
```
```
>>> not 'eFG%' in nba.columns
True
```
```
>>> not 'FT%' in nba.columns
True
```
```
>>> all(nba.iloc[4:8, 5] == [81, 38, 80, 81])
True
```
```
>>> all(nba.iloc[100:105, 23] == [2.1, 1.1, 1.3, 2.4, 1.5])
True
```


Tests for q2b (points: 4)
```
>>> nbaScaled.shape == (420, 18)
True
```


Tests for q2c (points: [1])
```
>>> len(sse) == 9
True
```


Tests for q2d (points: 8)
```
>>> clusterStats.shape == (4, 19)
True
```


Tests for q2e (points: [1])
```
>>> clusterStatsOrig.shape == (4, 18)
True
```


Tests for q3a (points: 4)
```
>>> music.shape == (100, 17)
True
```
```
>>> all(music.columns == ['Song', 'Artist', 'Streams (Billions)', 'Release Date', 'id', 'duration', 'energy', 'key', 'loudness', 'mode', 'speechiness', 'acousticness', 'instrumentalness', 'liveness', 'valence', 'tempo', 'danceability'])
True
```


Tests for q3b (points: 9)
```
>>> dg_single['leaves'][1:5] == [23, 11, 14, 21]
True
```
```
>>> dg_single['leaves'][20:24] == [8, 19, 25, 4]
True
```
```
>>> dg_single['leaves'][10:15] == [10, 13, 2, 5, 3]
True
```


Tests for q3c (points: 9)
```
>>> dg_complete['leaves'][1:5] == [1, 2, 3, 5]
True
```
```
>>> dg_complete['leaves'][10:15] == [27, 11, 13, 6, 20]
True
```
```
>>> dg_complete['leaves'][23:27] == [22, 16, 26, 7]
True
```


Tests for q3d (points: 9)
```
>>> dg_average['leaves'][1:5] == [15, 28, 23, 8]
True
```
```
>>> dg_average['leaves'][10:15] == [20, 3, 5, 13, 22]
True
```
```
>>> dg_average['leaves'][24:29] == [10, 11, 14, 27, 0]
True
```
